---
title       : Intro to Bash
description : Insert the chapter description here
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf

--- type:NormalExercise lang:python xp:100 skills:1 key:a38ebbbeb1
## Intro Bash

You can use the `env` command to display the current environment of the shell

*** =instructions
- Use `env` and play around in the console

*** =hint
- You don't have to program anything for the first instruction, just take a look at the first line of code.

*** =pre_exercise_code
```{python}
a=1337
```

*** =sample_code
```{python}
env
```

*** =solution
```{python}
env
```

*** =sct
```{python}

```
